```python
# Initialize Otter
import otter
grader = otter.Notebook("discussion.ipynb")
```

# Discussion 4

### Due Saturday April 23rd, 11:59:59PM


## Agenda: Plotting (Pandas, Seaborn, Matplotlib) and EDA

* **Pandas (plotting):** `.plot` dataFrame/Series methods conveniently plot tabular data (calls Matplotlib).
* **Seaborn:** A python plotting library similar to R's ggplot; makes common statistical plots easy.
* **Matplotlib:** Low-level plotting library built on Numpy; use it directly when you want to customize!


* **EDA:** Guide to exploring unknown data through visualizations - Univariate and Bivariate analysis.

```python
import pandas as pd
import seaborn as sns
import matplotlib
import matplotlib.pyplot as plt
import numpy as np
```

```python
from discussion import *
```

## 1. Plotting in `pandas` is as easy as `.plot()`

* `Series.plot()` plots a column.

```python
data = pd.read_csv('data/data.csv')
```

```python
data.head()
```

```python
# select a column from data
z0 = data['z0']
z0.head()
```

* Use a line plot to plot numeric data.
* `data.plot()` plots a line plot by default.
    - The x-axis is the index by default
    - Can be called out using the key-word argument `x`.

```python
# index is [0...1000]
z0.plot()
```

```python
# set index to plot correct x-axis
z0 = data.set_index('x').loc[:, 'z0']
z0.head()
```

```python
z0.plot()
```

```python
# set x-axis using a keyword argument
data.plot(x='x', y='z0')
```

### Plotting (quantitative) empirical distributions in Pandas

* Use the key-word argument `kind`
```
kind : str
    - 'hist' : histogram
    - 'box' : boxplot
    - 'kde' : Kernel Density Estimation plot
    ...
```
* The `hist` keyword by default uses 10 bins, and returns the *count* of observations within those bins.
    - use `density=True` to return a histogram whose area is normalized to 1.

```python
# histogram of z0 values; 
# 25 bins.
# density = normalized histogram

z0.plot(kind='hist', bins=25, density=True, ec='w')
```

```python
# kernel density estimate of the distribution
# smooth approximation of the empirical distribution

z0.plot(kind='kde')
```

```python
z0.plot(kind='box')
```

### Plotting (categorical) empirical distributions in Pandas

* Create a distribution from categorical columns using `value_counts`.
* Categorical columns should use *bar charts*.
* Use the key-word argument `kind`
```
kind : str
    - 'bar' : vertical bar plot
    - 'barh' : horizontal bar plot
    ...
```


```python
empdistr = data['id'].value_counts(normalize=True)
empdistr
```

```python
# nominal column
empdistr.plot(kind='bar')
```

```python
# ordinal column: the x-axis has a meaningful order
empdistr.sort_index().plot(kind='bar')
```

```python
# horizontal bar chart
empdistr.sort_index().plot(kind='barh')
```

### Plotting `pandas` DataFrames
* `DataFrame.plot()` plots the columns of a dataframe.
* Want multiple plot on the same axis? Get the data into the columns of a dataframe!

```python
data.set_index('x').head()
```

```python
# plot columns 'z0' and 'z1' with 'x' used as the x-axis
data.set_index('x')[['z0', 'z1']].plot()
```

```python
# plot columns 'z0' and 'z1' with 'x' used as the x-axis on seperate plots
data.set_index('x')[['z0', 'z1']].plot(subplots=True);
```

```python
# plot all columns using 'x' as x-axis; elongate plots with 'figsize' keyword
data.set_index('x').plot(subplots=True, figsize=(12,8));
```

### Scatter-plots with Pandas
* You can create scatter plots with `DataFrame.plot` by passing `kind='scatter'`. Scatter plot requires numeric columns for `x` and `y` axis. 
    * These can be specified by `x` and `y` keywords each.
* To plot multiple column groups in a single axes, repeat plot method specifying target `ax`. It is recommended to specify color and label keywords to distinguish each groups.

```python
data.plot(kind='scatter', x='z0', y='z1')
```

```python
# plot all the histograms and scatterplots in one plot!
# univariate + bivariate analysis
pd.plotting.scatter_matrix(data.drop(['id', 'x'], axis=1));
```

```python
df = pd.DataFrame(np.random.rand(50, 4), columns=['a', 'b', 'c', 'd'])

df.plot(kind='scatter', x='a', y='b');
```

There are other keywords that can be used with scatter. The keyword `c` may be given as the name of a column to provide colors for each point:

```python
samp = data.sample(100)
```

```python
samp.plot(kind='scatter', x='z0', y='z1', c='z3', s=50);
```

You can pass other keywords supported by matplotlib `scatter`. The example below shows a bubble chart using a dataframe `column` values as bubble size.

```python
samp.plot(kind='scatter', x='z0', y='z1', s=samp['x']);
```

## 2. Seaborn: pretty plotting made easy

To install `seaborn`, open a terminal and enter: 

On your laptop:
* `pip install seaborn==0.9`

or, if you are on a shared server (e.g. on `datahub.ucsd.edu`):

* `pip install --user seaborn==0.9`

The `seaborn` documentation has a *great* series of tutorials: https://seaborn.pydata.org/tutorial.html


```python
import seaborn as sns
sns.__version__
```

#### `sns.scatterplot`
* The relationship between `x` and `y` can be shown for different subsets of the data using the `hue`, `size`, and `style` parameters. 
* These parameters control what visual semantics are used to identify the different subsets. 
* It is possible to show up to three dimensions independently by using all three semantic types, but this style of plot can be hard to interpret and is often ineffective. 
    * Using redundant semantics (i.e. both `hue` and `style` for the same variable) can be helpful for making graphics more accessible.

Show a quantitative variable by using continuous colors:

```python
sns.scatterplot(data=data, x='z0', y='z1', hue='id')
```

Also show a quantitative variable by varying the size of the points:

```python
sns.scatterplot(data=data, x='z0', y='z1', size='id')
```

#### `sns.lmplot`

Plot a simple linear relationship between two variables:

```python
# plot a line of best fit
sns.lmplot(data=data, x='z0', y='z2');
```

#### `sns.distplot`

Plot the distribution with a histogram, kernel density estimate, and rug plot:

```python
z3 = data.sample(50)['z3']
sns.distplot(z3, hist=True, kde=True, rug=True)
```

#### `sns.boxplot`

Draw a vertical boxplot grouped by a categorical variable:

```python
sns.boxplot(data=data, x='id', y='z2')
```

## 3. Custom plots with `matplotlib`

* There are other great resources for learning the matplotlib API, for example, [this tutorial](https://www.southampton.ac.uk/~fangohr/training/python/notebooks/Matplotlib.html)
* A lot of examples and tutorials are also given in the [official website](https://matplotlib.org/stable/gallery/index.html)

```python
import matplotlib.pyplot as plt
```

### Matplotlib `axes` objects and Pandas plots

* An 'Axes' object contains the elements of a single plot.
    - contains a coordinate system (axis elements), 
    - the plot elements (e.g. line, bar), 
    - labels, 
    - tick-marks, etc.
    
* A `DataFrame.plot()` method call returns an `axes` object

#### Anatomy of a Matplotlib Figure
[Source](https://matplotlib.org/3.5.0/gallery/showcase/anatomy.html)

Understanding the terminology and anatomy is really helpful in efficient searching and customization of matplotlib plots.

<img src="imgs/anatomy_fig_matplotlib.png" width="60%"/>

```python
# notice the <matplotlib.axes._subplots.AxesSubplot at 0x1a21f7bcf8>
data.set_index('x')['z0'].plot()
```

```python
# save the plot as an variable
ax = data.set_index('x')['z0'].plot()
```

```python
# get name of x-axis
ax.get_xlabel()
```

```python
# get y-axis tick-labels
list(ax.get_yaxis().get_majorticklabels())
```

```python
ax = data.set_index('x')['z0'].plot()
ax.set_xlabel('hi, this is my new axis label!')
ax.set_title('hi this is my new title!');
```

#### You can add elements to an Axes object

* The Pandas `.plot` method can add a plot to an existing Axes object using the `ax` keyword

```python
ax = data['z0'].plot()

# add z1 to Axes
data['z1'].plot(ax=ax)

# add a vertical line using matplotlib
plt.plot([40,40],[-400, 300])

# add a point using matplotlib
plt.plot(15,-200, marker='x', markersize=10, color='red')
```

#### You can add a scatterplot to an existing scatterplot

```python
ax = data.plot(kind='scatter', x='z0', y='z1', alpha=0.3)

# the 'ax' keyword in Pandas plot method attaches the new plot to an existing Axes object
data.plot(kind='scatter', x='z0', y='z3', ax=ax, c='g', alpha=0.3)
```

### Matplotlib `figure` and adding to empty subplots

* A 'Figure' object is a top-level container for all plotting objects.
    - controls overall size, title, fonts, coordination between different elements of subplots.

<img src="https://i.stack.imgur.com/HZWkV.png" width="25%">  

* Instantiate an empty figure containing multiple plots with `plt.subplots`
    - `fig, axes = plt.subplots(R, C)` returns a figure `fig` and an multi-array of `axes`.
    - `axes` has `R` rows and `C` columns corresponding to the subplots laid out on a grid.
    - The `axes` are initially empty; they need to be given data to plot.
   

```python
fig, axes = plt.subplots(1, 2)
```

```python
len(axes), type(axes), type(axes[0])
```

```python
fig, axes = plt.subplots(1, 2, figsize=(12,4))

df = data.set_index('x')
df['z0'].plot(ax=axes[0], title='z0')
df['z1'].plot(ax=axes[1], title='z1')
```

```python
fig, axes = plt.subplots(2, 1, sharex=True)

df = data.set_index('x')
df['z0'].plot(ax=axes[0], title='z0')
df['z1'].plot(ax=axes[1], title='z1')
```

### Practice: plots and groupby

* Can we plot histograms of `z2` for each value of `id`?

```python
# Hard to understand!
data.drop('x', axis=1).groupby('id')['z2'].plot(kind='hist', alpha=0.3);
```

```python
data['id'].nunique()
```

```python
grps = data.groupby('id')
for k, gp in grps:
    print('**** ' + str(k) + ' ****', grps.get_group(k).head().to_string(), sep='\n', end='\n\n')
```

```python
fig, axes = plt.subplots(2, 2, sharex=True, sharey=True)

for k, gp in data.groupby('id'):
    x_idx = k // 2
    y_idx = k % 2
    ax = axes[x_idx, y_idx]
    title = 'id = %d' % k
    gp['z2'].plot(kind='hist', density=True, ax=ax, title=title, ec='w')
    
fig.suptitle('Distribution of z2 by id-number');

```

**Try out (Optional)**: Can you plot the histograms of each column by `id`? Each row should contain the histograms by `id` of a single variable (there should be 3 rows and 4 columns). Write this generally enough to handle an arbitrary number of variables and values of `id`.

```python

```

### Exploring Unknown Data through Visualizations

- Univariate Analysis
    - Numerical/Continuous variables - Histogram, density plot, box plot etc.
    - Categorical variables - Bar plot of counts


- Bivariate Analysis
    - Continuous vs. Continuous - What kind of plot is best?
    - Continuous vs. Categorical - What kind of plot is best?
    - Categorical vs. Categorical - What kind of plot is best?
    
Try to use this as a starting point in your EDA on an unknown dataset. These univariate and bivariate analysis should generate more hypothesis and questions to dig further.

Practice this on the below `tips` dataset.

### Practice problems

* Below is a dataset in the seaborn package that contains data on restaurant bills and (service) tips.
* Try to understand the dataset via plotting using the examples in the notebook.
    - Plot histograms and boxplots for quantitative columns
    - Plot counts of categorical values using bar plots
    - Plot a scatter plot of `tip` vs `total_bill` -- is the relationship linear?

```python
tips = sns.load_dataset('tips')
```

```python
tips.head()
```

```python
# ======== Univariate ===========

# sns.histplot(data=tips, x= 'total_bill', ec='w')
# sns.boxplot(data=tips, x= 'total_bill')

# sns.countplot(data=tips, x = 'sex')

# ======== Bivariate ============

# sns.scatterplot(data=tips, x='total_bill', y='tip')

# sns.boxplot(data=tips, x= 'sex', y='total_bill')

# cross_tab_prop = pd.crosstab(index=tips['sex'], columns=tips['smoker'], normalize="index")
# cross_tab_prop
# cross_tab_prop.plot(kind='bar', stacked=True, figsize=(10, 6))
# plt.legend(loc="upper left", ncol=2)
# plt.xlabel("Sex")
# plt.ylabel("Proportion")
# plt.show()
```

```python

```

**Question 1**

Plot the counts of meals in `tips` by day. Your plotting function, `plot_meal_by_day` should return an `matplotlib.axes._subplots.AxesSubplot` object; your plot should look like the plot below. 

Note: You don't need to have exact same colors, but the plot orientation, axis order, and title should match.

<img src="imgs/barh.png" width="50%"/>

```python

```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
tips = sns.load_dataset('tips')
q1_fig = plot_meal_by_day(tips)
```

```python
grader.check("q1")
```

**Question 2**

Plot a seaborn scatterplot using the `tips` data by day. Your plotting function, `plot_bill_by_tip` should return a `matplotlib.axes._subplots.AxesSubplot` object; your plot should look like the plot below.
* `tip` is on the x-axis.
* `total_bill` is on the y-axis.
* color of the dots are given by `day`.
* size of the dots are given by `size` of the table.

<img src="imgs/scatter.png" width="50%"/>

```python

```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
tips = sns.load_dataset('tips')
q2_fig = plot_bill_by_tip(tips)
```

```python
grader.check("q2")
```

**Question 3**

Plot a figure with two subplots side-by-side. The left plot should contain the **counts** of tips *as a percentage of the total bill*. The right plot should contain the **density plot** of tips as a percentage of the total bill. Your plotting function, `plot_tip_percentages` should return a `matplotlib.Figure` object; your plot should look like the plot below (use 10 bins).

<img src="imgs/hist.png" width="50%"/>

```python

```

```python

```

```python
fig, axes = plt.subplots(1, 2)

# plot axes[0]
...
# plot axes[1]
...
# add the title to fig
...
```

```python
# don't change this cell -- it is needed for the tests to work
tips = sns.load_dataset('tips')
q3_fig = plot_tip_percentages(tips)
```

```python
grader.check("q3")
```

## Congratulations! You're done!

* Submit your `.py` file to Gradescope. Note that you only need to submit the `.py` file; this notebook should not be uploaded. Make sure that all of your work is in the `.py` file and not here by running the doctests: `python -m doctest discussion.py`.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: None)
```
>>> isinstance(q1_fig, matplotlib.figure.Figure)
True
```
failure_message: doctest
```
>>> q1_fig.get_axes()[0].get_title()
'Counts of meals by day'
```
failure_message: need correct title
```
>>> [str(x.get_text()) for x in q1_fig.get_axes()[0].get_yticklabels()]
['Thur', 'Fri', 'Sat', 'Sun']
```
failure_message: y-axis sorted and labeled correctly?


Tests for q2 (points: None)
```
>>> isinstance(q2_fig, matplotlib.figure.Figure)
True
```
failure_message: doctest
```
>>> q2_fig.get_axes()[0].get_xlabel()
'tip'
```
failure_message: x-axis label incorrect
```
>>> q2_fig.get_axes()[0].get_ylabel()
'total_bill'
```
failure_message: y-axis label incorrect


Tests for q3 (points: None)
```
>>> isinstance(q3_fig, matplotlib.figure.Figure)
True
```
failure_message: doctest
```
>>> len(q3_fig.get_axes()) == 2
True
```
failure_message: need two plots (axes)
```
>>> q3_fig.get_axes()[0].get_title()
'counts'
```
failure_message: left plot does not have correct title
```
>>> q3_fig.get_axes()[1].get_title()
'normalized'
```
failure_message: right plot does not have correct title
