```python
# Initialize Otter
import otter
grader = otter.Notebook("discussion.ipynb")
```

# Discussion 7

### Due Saturday May 7th, 11:59:59PM


# Web Scraping

```python
import os
import numpy as np
import pandas as pd
import requests
import time
from bs4 import BeautifulSoup
import re
```

```python
from discussion import *
```

## Let's scrape some books data!
* We will be scraping some book information from http://books.toscrape.com/index.html
* Note that this website is built for scraping purposes

### 1. HTML

* **Hyper Text Markup Language** is used to provide content to a web page and instruct web browsers (like Chrome, Safari etc.) on how to structure that content
* HTML tags form the builing blocks of structuring a webpage - [HTML Tags Reference](https://www.w3schools.com/tags/default.asp)
* **CSS** is used to style the web pages and can be used inplace inside the HTML code or by using external css styling files
* **JavaScript** is used to provide interactivity within the browser


* Note that you don't need be a web designer or developer to scrape data from standard websites. Tools like BeautifulSoup with basic knowledge of HTML tags will provide a great starting point

### 2. Get all the book urls from single catalogue page

```python
url = 'http://books.toscrape.com/catalogue/page-1.html'
```

```python
book_request = requests.get(url)
soup = BeautifulSoup(book_request.text)
soup
```

The key to scraping is in identifying the container/block/division that is unique to your requirement. Also note that there might exist multiple ways to scrape the same information.

<img src='imgs/scraping_container.png'>

```python
articles = soup.find_all('article', {"class": "product_pod"})
articles
```

```python
articles[0].find('a')
```

```python
book_links = [art.find('a').get('href') for art in articles]
book_links
```

### 3. Parse Book information from a book URL

Let's try and extract the title, price, instock availability and rating of the book.

```python
# For the first book
book_url = 'http://books.toscrape.com/catalogue/' + book_links[0]
```

```python
book_soup = BeautifulSoup(requests.get(book_url).text)
book_soup
```

#### 3.1 Book Title

```python
book_soup.find_all('h1')
# book_soup.find_all('h2')
```

```python
name = book_soup.find('article', {'class':'product_page'}).find('h1').text
# name = book_soup.find('article', {'class':'product_page'}).h1.text
name
```

#### 3.2 Book Price

```python
price = book_soup.find('p', {'class':'price_color'}).text
price
```

```python
price = float(price[2:])
price
```

#### 3.3 Book Availability

```python
instock = book_soup.find('p', {'class':'instock availability'}).text
instock
```

```python
instock_units = re.findall(r'\d+', instock)
# Note that you can do this using string methods like .isdigit() as well
instock_units
```

#### 3.4 Book Rating

```python
rating = book_soup.find('p', class_ = re.compile('star-rating')).get('class')[1]
rating
```

**Question 1**

Get all the catalog urls of the website in the form of a list. Notice that each page of the website has a very specific pattern.

Hint: Try clicking the "next" button at the bottom of the page to see how the URL changes.

```python

```

```python

```

```python
grader.check("q1")
```

**Question 2**

Extract all the category names (`Travel`, `Mystery`..) on the home page (http://books.toscrape.com/index.html) and return them as a list. Notice to remove all the extra whitespaces if any.

```python

```

```python

```

```python
# don't change this cell -- it is needed for the tests to work
cats = book_categories()
```

```python
grader.check("q2")
```

## Congratulations! You're done!

* Submit your `.py` file to Gradescope. Note that you only need to submit the `.py` file; this notebook should not be uploaded. Make sure that all of your work is in the `.py` file and not here by running the doctests: `python -m doctest discussion.py`.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: None)
```
>>> urls = get_website_urls()
>>> len(urls)==50
True
```
failure_message: doctest
```
>>> urls = get_website_urls()
>>> page_nums = [int(re.findall('page-(\d)',url)[0]) for url in urls]
>>> sum(page_nums)==150
True
```
failure_message: sum of all page numbers


Tests for q2 (points: None)
```
>>> len(cats)==50
True
```
failure_message: doctest
```
>>> 'Classics' in cats
True
```
failure_message: doctest
```
>>> min(cats)=='Academic' and max(cats)=='Young Adult'
True
```
failure_message: max and min by alphabet sort
